/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion libraryDefinition:
 *   libraryName importOrExport* partDirective* topLevelDefinition*
 * ;
 * @description Checks that only the library name is required for a library
 * and other directives and definitions can be omitted.
 * @author vasya
 * @reviewer hlodvig
 * @reviewer msyabro
 * @reviewer rodionov
 */

import "13_Libraries_and_Scripts_A04_t02_lib.dart";

main() {
}
