/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion libraryDefinition:
 *   libraryName import* partDirective* topLevelDefinition*
 * ;
 * @description Checks that a library is parsed without errors
 * if all allowed directives are present in permitted numbers and in the correct order.
 * @author msyabro
 * @reviewer rodionov
 */

import "13_Libraries_and_Scripts_A04_t20_lib.dart";

main() {
}
