/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion libraryName:
 *    metadata library identiﬁer (‘.’ identiﬁer)* ‘;’
 * ;
 * An explicitly named library begins with the word library (possibly prefaced
 * with any applicable metadata annotations), followed by a qualiﬁed identiﬁer
 * that gives the name of the library.
 * @description Checks that valid library directives (according to the spec)
 * do not cause compile-time errors.
 * @author msyabro
 * @reviewer rodionov
 * @issue 5470
 */

import "13_Libraries_and_Scripts_A03_t01_p1.dart";
import "13_Libraries_and_Scripts_A03_t01_p2.dart";
import "13_Libraries_and_Scripts_A03_t01_p3.dart";
//import "13_Libraries_and_Scripts_A03_t01_p4.dart";
//import "13_Libraries_and_Scripts_A03_t01_p5.dart";
//import "13_Libraries_and_Scripts_A03_t01_p6.dart";

main() {
}
