/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion It is a run time error if library does not declare or import a top level function main().
 * @description Checks that if script does not have a top-level function main()
 * it can be imported.
 * @author vasya
 * @reviewer msyabro
 */

#import("4_Library4.dart");

