/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion WHITESPACE:
 * ('\t' | ' ' | NEWLINE)+
 * ;.
 * @description Checks that all of these characters are indeed considered
 * whitespace by verifying that all and any of their occurrences within
 * grammatical productions are ignored. 
 * @author rodionov
 * @reviewer iefremov
 */
import "../../../../Utils/expect.dart";

main() {
  Expect. 	
  	 		
  equals( 	 
  		3, 1  
  		+
  		
  		
   	 2);
}
