/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion Static type annotations are used during static checking and when
 * running programs in checked mode. They have no effect whatsoever in production mode.
 * @description TODO:
 * @author iefremov
 * @needsreview TODO: invent a way to check that static type annotations do not cause compile-time errors in scripting mode.
 * It looks like it should be supported by harness.
 */

main() {

}
