// Copyright (c) 2024, the Dart project authors.  Please see the AUTHORS file
// for details. All rights reserved. Use of this source code is governed by a
// BSD-style license that can be found in the LICENSE file.

/// @assertion Library file to test identifier `show`.
///
/// @description Library file to test identifier `show`.
/// @author sgrekhov22@gmail.com

// SharedOptions=--enable-experiment=unquoted-imports

String libraryId = "show";
