// Copyright (c) 2021, the Dart project authors.  Please see the AUTHORS file
// for details. All rights reserved. Use of this source code is governed by a
// BSD-style license that can be found in the LICENSE file.

/// @assertion A script is a library whose exported namespace (19.2) includes a
/// top-level function declaration named main that has either zero, one or two
/// required arguments.
///
/// @description Checks main function can have zero required arguments
/// @author iarkh@unipro.ru

main() {}
