/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion Returns hash code of the string
 * @description Check that hash code is calculated correctly
 * @author msyabro
 * @needsreview calculating method is undocumented
 */


main() {
}
