/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion static num pow(num x, num exponent)
 * @description Checks set of values.
 * @author msyabro
 * @reviewer pagolubev
 * @3rdparty sputnik-v1: S15.8.2.13_A24.js
 * @needsreview Returned value precision is not specified.
 * Checks exact match with values calculated with LibC.
 */

main() {
  int vnum = 64;
  var x1 = new List(64);
  x1[0] = 0.00000000000000000000;
  x1[1] = 0.25396825396825395000;
  x1[2] = 0.50793650793650791000;
  x1[3] = 0.76190476190476186000;
  x1[4] = 1.01587301587301580000;
  x1[5] = 1.26984126984126980000;
  x1[6] = 1.52380952380952370000;
  x1[7] = 1.77777777777777770000;
  x1[8] = 2.03174603174603160000;
  x1[9] = 2.28571428571428560000;
  x1[10] = 2.53968253968253950000;
  x1[11] = 2.79365079365079350000;
  x1[12] = 3.04761904761904740000;
  x1[13] = 3.30158730158730140000;
  x1[14] = 3.55555555555555540000;
  x1[15] = 3.80952380952380930000;
  x1[16] = 4.06349206349206330000;
  x1[17] = 4.31746031746031720000;
  x1[18] = 4.57142857142857120000;
  x1[19] = 4.82539682539682510000;
  x1[20] = 5.07936507936507910000;
  x1[21] = 5.33333333333333300000;
  x1[22] = 5.58730158730158700000;
  x1[23] = 5.84126984126984090000;
  x1[24] = 6.09523809523809490000;
  x1[25] = 6.34920634920634890000;
  x1[26] = 6.60317460317460280000;
  x1[27] = 6.85714285714285680000;
  x1[28] = 7.11111111111111070000;
  x1[29] = 7.36507936507936470000;
  x1[30] = 7.61904761904761860000;
  x1[31] = 7.87301587301587260000;
  x1[32] = 8.12698412698412650000;
  x1[33] = 8.38095238095238140000;
  x1[34] = 8.63492063492063440000;
  x1[35] = 8.88888888888888930000;
  x1[36] = 9.14285714285714230000;
  x1[37] = 9.39682539682539720000;
  x1[38] = 9.65079365079365030000;
  x1[39] = 9.90476190476190510000;
  x1[40] = 10.15873015873015800000;
  x1[41] = 10.41269841269841300000;
  x1[42] = 10.66666666666666600000;
  x1[43] = 10.92063492063492100000;
  x1[44] = 11.17460317460317400000;
  x1[45] = 11.42857142857142900000;
  x1[46] = 11.68253968253968200000;
  x1[47] = 11.93650793650793700000;
  x1[48] = 12.19047619047619000000;
  x1[49] = 12.44444444444444500000;
  x1[50] = 12.69841269841269800000;
  x1[51] = 12.95238095238095300000;
  x1[52] = 13.20634920634920600000;
  x1[53] = 13.46031746031746000000;
  x1[54] = 13.71428571428571400000;
  x1[55] = 13.96825396825396800000;
  x1[56] = 14.22222222222222100000;
  x1[57] = 14.47619047619047600000;
  x1[58] = 14.73015873015872900000;
  x1[59] = 14.98412698412698400000;
  x1[60] = 15.23809523809523700000;
  x1[61] = 15.49206349206349200000;
  x1[62] = 15.74603174603174500000;
  x1[63] = 16.00000000000000000000;

  var x2 = new List(64);
  x2[0] = -16.00000000000000000000;
  x2[1] = -15.49206349206349200000;
  x2[2] = -14.98412698412698400000;
  x2[3] = -14.47619047619047600000;
  x2[4] = -13.96825396825396800000;
  x2[5] = -13.46031746031746000000;
  x2[6] = -12.95238095238095300000;
  x2[7] = -12.44444444444444500000;
  x2[8] = -11.93650793650793700000;
  x2[9] = -11.42857142857142900000;
  x2[10] = -10.92063492063492100000;
  x2[11] = -10.41269841269841300000;
  x2[12] = -9.90476190476190510000;
  x2[13] = -9.39682539682539720000;
  x2[14] = -8.88888888888888930000;
  x2[15] = -8.38095238095238140000;
  x2[16] = -7.87301587301587350000;
  x2[17] = -7.36507936507936560000;
  x2[18] = -6.85714285714285770000;
  x2[19] = -6.34920634920634970000;
  x2[20] = -5.84126984126984180000;
  x2[21] = -5.33333333333333390000;
  x2[22] = -4.82539682539682600000;
  x2[23] = -4.31746031746031810000;
  x2[24] = -3.80952380952381020000;
  x2[25] = -3.30158730158730230000;
  x2[26] = -2.79365079365079440000;
  x2[27] = -2.28571428571428650000;
  x2[28] = -1.77777777777777860000;
  x2[29] = -1.26984126984127070000;
  x2[30] = -0.76190476190476275000;
  x2[31] = -0.25396825396825484000;
  x2[32] = 0.25396825396825307000;
  x2[33] = 0.76190476190476275000;
  x2[34] = 1.26984126
