/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion const FallThroughError()
 * @description Checks that this method causes no error and
 * @author rodionov
 * @needsreview Undocumented
 * @needsreview issue 1269
 */
 
main() {
  var fte = const FallThroughError();
}
