/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion factory StreamController.broadcast({void onListen(), void onCancel(), bool sync: false})
 * It is not allowed to call add, addError, or close before a previous call has returned.
 * @description TODO invent how to check
 */

//import "dart:async";
//import "../../../Utils/async_utils.dart";
//import "../../../Utils/expect.dart";

main() {
}
