/*
 * Copyright (c) 2014, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion final String localName
 * The local part of the qualified name of an element.
 * @description Checks expected attribute values.
 * @note tbd
 */
import "dart:html";
import "../../../Utils/expect.dart";
import "../testcommon.dart";

main() {
}
