/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion An unqualified assignment is of the form identifier = rhs.
 * Where:
 * identifier is a simple identifier.
 * rhs is expression.
 * @description Check that unqualified assignment is a valid expression
 * @author msyabro
 */


int x;

main() {
  x = 1;
}
