/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion const IntegerDivisionByZeroException()
 * @description Checks that this constructor executes without error.
 * @author rodionov
 */
 
main() {
  new IntegerDivisionByZeroException();
}
