/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion Every class has a single superclass except class Object which has no superclass. 
 * @description Uncheckable as there're no reflection facilities in Dart at this point.
 * @author vasya
 * @reviewer rodionov
 */

main() {
}

