/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion Multiline comments openers "/*" must be ignored inside multiline comments.
 * @description Try to place "/*" inside multiline comment block
 * @author akuznecov
 */

/*
/**/

void main() {}