// Copyright (c) 2024, the Dart project authors.  Please see the AUTHORS file
// for details. All rights reserved. Use of this source code is governed by a
// BSD-style license that can be found in the LICENSE file.

/// @assertion It is a compile-time error if a library contains two part
/// directives with the same URI.
///
/// @description Checks that it is a compile-time error if a library contains
/// two part directives with the same URI.
/// @author sgrekhov22@gmail.com

part of 'uri_t01.dart';

var uri_t01_part_var = "uri_t01_part_var";
