/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion libraryName:
 *   '#' 'library' '(' stringLiteral ')' ';'
 * ;
 * @description Checks that valid library directives (according to the spec)
 * do not cause compile-time errors.
 * @author msyabro
 * @reviewer rodionov
 * @note http://code.google.com/p/dart/issues/detail?id=1712
 */

#import("12_Libraries_and_Scripts_A03_t01_p1.dart");
#import("12_Libraries_and_Scripts_A03_t01_p2.dart");
#import("12_Libraries_and_Scripts_A03_t01_p3.dart");
#import("12_Libraries_and_Scripts_A03_t01_p4.dart");
#import("12_Libraries_and_Scripts_A03_t01_p5.dart");
#import("12_Libraries_and_Scripts_A03_t01_p6.dart");

// no longer valid as of 0.08
//#import("12_Libraries_and_Scripts_A03_t01_p7.dart");

main() {
}
