/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion It is a compile-time error to import into a library two or more namespaces
 * that define the same name.
 * @description This assertion seems to duplicate the one tested in A03, so no new tests are needed.
 * @author rodionov
 */

main() {
}
