/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion const DeferredLibrary(String libraryName, {String uri})
 * @description Checks that DeferredLibrary loads in lazy way.
 * @author kaigorodov
 */

library DeferredLibrary_A01_t01.lib;

String method() {
   return "DeferredLibrary_A01_t01.lib";
}
