/*
 * Copyright (c) 2011, the Dart project authors.  Please see the AUTHORS file
 * for details. All rights reserved. Use of this source code is governed by a
 * BSD-style license that can be found in the LICENSE file.
 */
/**
 * @assertion The logical-or expression is of the form: 
 * expression1 || expression2
 * Where expression1 and expression2 are expressions in condition-expression context.
 * @description Check if operator || is a valid operator
 * @author msyabro
 * @reviewer akuznecov
 */
 

void main() {
  true || false;
}
